---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Part 1
challengeType: 0
dashedName: part-1
---

# --description--

Bem vindo ao projeto painel de controle! Você irá usar a biblioteca de visualização de dados do JavaScript, D3, para construir a visualização dos seus seguidores de redes sociais. Que consiste de um gráfico linear, diagrama de pizza e legendas.

Primeiramente, você deverá criar o arquivo HTML. Comece adicionando a declaração `<!DOCTYPE html>` no topo do seu arquivo para dizer ao navegador que tipo de documento se trata.

# --hints--

texto de teste

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
